<a href="https://colab.research.google.com/github/GGuerra0301/GGuerra0301/blob/master/Introductionto_Python.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "GABRIELA GUERRA"
COLLABORATORS = ""

```

---

# Objective

The objective of this notebook is to explore how to use notebooks for homework.
You will be submitting your homework during the course using these notebooks.

```python
import numpy as np
```

<!-- grade_id: grade_vowels, points: 0, grade: false, solution: true -->
```python
def remove_vowels(s: str) -> str:
    """
    This function removes all vowels from strings
    Example: "hello" -> "hll"
    """
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: grade_remove_vowels, points: 1, grade: true, solution: false -->
```python
assert remove_vowels("hola") == "hl"
assert remove_vowels("this was a good day") == "ths ws  gd dy"
```

Output:
```
<ipython-input-4-78978080ccd6> in <module>
----> 1 assert remove_vowels("hola") == "hl"
      2 assert remove_vowels("this was a good day") == "ths ws  gd dy"

<ipython-input-3-806a8c701360> in remove_vowels(s)
      6 
      7     # YOUR CODE HERE
----> 8     raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: sum_of_list, points: 0, grade: false, solution: true -->
```python
def sum_list_of_floats(a: list, b:list) -> list:
    """
    Returns the coordinate wise sum of a and b.
    """
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: grade_sum, points: 1, grade: true, solution: false -->
```python
assert np.allclose(
    sum_list_of_floats(
        [1, 2, 3], [4, 5, -2.3]),
    [5, 7, 0.7]
)
assert np.allclose(
    sum_list_of_floats(
        [0, 0], [1, -1]),
    [1, -1]
)
```